rule_files:
    - cluster_rules.yaml

evaluation_interval: 1m

tests:
    - interval: 1m
      input_series:
          - series: 'velero_backup_failure_total{container="velero", endpoint="http-monitoring", instance="111.111.111.11:8085", job="velero", namespace="velero", pod="velero", schedule="velero-fusion-backup", service="velero"}'
            values: '0+0x24 1+0x10 2 2 2'  #'0-24=0 25-35=1 36=2 37=2 38=2' note: the alert has for: 15s

      # Unit test VeleroBackupFailure Alert.
      alert_rule_test:
          - eval_time: 10m
            alertname: VeleroBackupFailure
            exp_alerts: []
          - eval_time: 25m
            alertname: VeleroBackupFailure
            exp_alerts: []
          - eval_time: 26m
            alertname: VeleroBackupFailure
            exp_alerts:
                - exp_labels:
                      severity: P3-Minor
                      scope: cluster
                      resolves: never
                      schedule: velero-fusion-backup
                  exp_annotations: 
                      message: Failed backup in Velero schedule velero-fusion-backup.
          - eval_time: 27m
            alertname: VeleroBackupFailure
            exp_alerts:
                - exp_labels:
                      severity: P3-Minor
                      scope: cluster
                      resolves: never
                      schedule: velero-fusion-backup
                  exp_annotations: 
                      message: Failed backup in Velero schedule velero-fusion-backup.
          - eval_time: 34m
            alertname: VeleroBackupFailure
            exp_alerts:
                - exp_labels:
                      severity: P3-Minor
                      scope: cluster
                      resolves: never
                      schedule: velero-fusion-backup
                  exp_annotations: 
                      message: Failed backup in Velero schedule velero-fusion-backup.
          - eval_time: 35m
            alertname: VeleroBackupFailure
            exp_alerts: []
          - eval_time: 37m
            alertname: VeleroBackupFailure
            exp_alerts:
                - exp_labels:
                      severity: P3-Minor
                      scope: cluster
                      resolves: never
                      schedule: velero-fusion-backup
                  exp_annotations: 
                      message: Failed backup in Velero schedule velero-fusion-backup.


        


